// Numbas version: finer_feedback_settings
{"name": "Matrix linear combination", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Matrix linear combination", "tags": [], "metadata": {"description": "", "licence": "All rights reserved"}, "statement": "", "advice": "", "rulesets": {}, "extensions": [], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "matrix(map(map(random(-1..4),x,1..n),y,1..m))", "description": "", "templateType": "anything", "can_override": false}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "matrix(map(map(random(-2..3),x,1..n),y,1..m))", "description": "", "templateType": "anything", "can_override": false}, "m": {"name": "m", "group": "Ungrouped variables", "definition": "random(2..3)", "description": "<p>Matrix number of rows, columns will be 5-rows, so 2x3 or 3x2.</p>", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Ungrouped variables", "definition": "5-m", "description": "<p>Dimensions of matrix are 3x2 or 2x3</p>", "templateType": "anything", "can_override": false}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "2*a-3*b", "description": "<p>2A-3B</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a", "b", "m", "n", "c"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "matrix", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Consider the following two matrices:</p>\n<p>\\[ A = \\var{A} \\text{ and } B=\\var{B} \\]</p>\n<p>Calculate the linear combination \\(2A-3B\\).</p>\n<p>Type your answer below (you should choose the shape first).</p>\n<p>If you do not get the right answer, feel free to click <em>Reveal Answers</em> to see which elements you got wrong.</p>", "correctAnswer": "{C}", "correctAnswerFractions": false, "numRows": "3", "numColumns": "3", "allowResize": true, "tolerance": 0, "markPerCell": false, "allowFractions": false, "minColumns": 1, "maxColumns": 0, "minRows": 1, "maxRows": 0}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "type": "question", "contributors": [{"name": "David Hodge", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/4151/"}], "resources": []}]}], "contributors": [{"name": "David Hodge", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/4151/"}]}